# Poker Dice

## Game mechanics

In the game of Poker Dice, two players each roll a hand of five standard (cubical or d6) dice. The goal is to have the highest ranking hand. Hand ranks are as follows, from low to high.

0. Nothing
2. One pair
3. Two pair (but not a full house)
4. Three of a kind (without an accompanying pair)
5. Short straight (either 1, 2, 3, 4 or 2, 3, 4, 5 or 3, 4, 5, 6 (other value irrelevant))
6. Straight (either 1, 2, 3, 4, 5 or 2, 3, 4, 5, 6)
7. Full house (three of a kind plus one pair)
8. Four of a kind (other value irrelevant)
9. Five of a kind

After the initial roll, each player may reroll *any of the dice* in their hand (whichever ones they choose). There are two rounds of rerolls and then the highest hand wins.

## Setup


<!-- grade_id: 89aa20, points: 1, grade: true, solution: false -->
```python
import os.path
assert(os.path.isfile("PokerDiceHand.py"))

from PokerDiceHand import MSDie
from PokerDiceHand import Hand
from PokerDiceHand import PokerDiceHand

MSDie(6) # Constructor will raise exception if class is ill-defined
Hand(6, 6)
PokerDiceHand()
```

<!-- grade_id: d6358a, points: 0, grade: false, solution: true -->
```python
class Player(PokerDiceHand):
    def __init__(self, name):
        """
        Initializes the player's PokerDiceHand (self.hand) and name (self.name).
        """
        
        self.hand = PokerDiceHand()
        self.name = name
        
    def __str__(self):
        """
        Return a string representation of the player's hand. Specifically, returns
        a string containing an announcement, followed by the 5 values, separated 
        by spaces: "Player 1's hand: 1 3 5 4 2", e.g. Use the player's name in the
        announcement.
        """

        return f"{self.name}'s hand: {self.hand}"
        
    def prompt_for_reroll(self):
        """
        Prompt the user for input to determine if a reroll will occur. The prompt
        string should look like:
        "Player 1, enter the dice to reroll (e.g., '134') or p to pass: "
        Make sure to use the player's name and to leave one space inside the string
        after the `:`. 
        
        Return a list of the values the player entered, but as `int`s, not `str`s.
        If the player chooses `p` to pass, return None.
        """

        choice = input(f"{self.name}, enter the dice to reroll (e.g., '134') or p to pass: ")
        list(choice)
        lst = []
        if choice.lower() != "p":
            for i in choice:
                lst.append(int(i))
        return lst
#            return [int(i) for i in choice]
```

<!-- grade_id: d90c36, points: 1, grade: true, solution: false -->
```python
# Check that __init__ does what it is supposed to do.
from nose.tools import assert_equal
p = Player("Philip J. Fry")
assert_equal(type(p.hand), type(PokerDiceHand()))
assert_equal(p.name, "Philip J. Fry")

# Check that show_hand does what it is supposed to do.
dice = [MSDie(6) for _ in range(5)]
values = [5, 6, 2, 3, 3]
for i, x in enumerate(dice):
    x.current_value = values[i]
p.hand.dice = dice
assert_equal(str(p), f"Philip J. Fry's hand: {p.hand}")
```

<!-- grade_id: a39a42, points: 0, grade: false, solution: true -->
```python
class Game():

    def __init__(self):
        """
        Setup a Game object. The Game object has just one attribute,
        self.players. This is a list containing two Player objects, representing
        the two players of the game. We could just have the players separately,
        but in coding the rest of the class, it's convenient to have a list. Name
        the two players "Player 1" and "Player 2".
        
        Notice that when the Player objects are created, the Player class's constructor
        (which was already written in the cell above) generates their PokerDiceHands, so
        we don't need to do that here.
        """

        self.players = [Player("Player 1"), Player("Player 2")]

    def play(self):
        """

        This method manages the game: printing appropriate messages to the
        screen, keeping track of rerolls (which it receives from the Player objects),
        delegating the rerolls to the PokerDiceHand objects, fetching the ranks of
        the PokerDiceHand objects, comparing them, and announcing the results. It 
        should allow the users to play more games indefinitely, so use a `while` loop 
        to keep starting new games until the answer to the "Play again?" prompt is 'n'.
        (Consult the sample output for reference.)
        """

        replay = "y"
        while replay == "y":
            self.players[0].hand.roll()
            self.players[1].hand.roll()
            for x in range(2):
                for t_player in self.players:
                    #print(t_player.prompt_for_reroll())
                    print(f"{t_player.name} roll: {t_player.hand}")
                    t_player.hand.reroll(t_player.prompt_for_reroll())
                    print(f"{t_player.name} after reroll: {t_player.hand}")
            rank_1 = self.players[0].hand.rank()
            rank_2 = self.players[1].hand.rank()
            print(f"{self.players[0]} {self.players[0].hand.rank()}")
            print(f"{self.players[1]} {self.players[1].hand.rank()}")
            order = ["has nothing", "has one pair", "has two pairs", "has three of a kind", "has short straight", "has straight", "has full house", "has four of a kind", "has five of a kind"]
            if order.index(rank_1) > order.index(rank_2):
                print(f"{self.players[0]} WIN!")
            elif order.index(rank_1) < order.index(rank_2):
                print(f"{self.players[1]} WIN!")
            else:
                print("It is a tie!")
            replay = input("Play again? y/n")
        print("GAME OVER!")
```

<!-- grade_id: 5eefb8, points: 1, grade: true, solution: false -->
```python
# Check that __init__ does what it is supposed to
from nose.tools import assert_equal
g = Game()
assert_equal(g.players[0].name, "Player 1")
assert_equal(g.players[1].name, "Player 2")
```

## Time to play!


```python
g = Game()
g.play()
```